rule_files:
  - prometheus_rules.yaml
evaluation_interval: 1m
tests:
  - interval: 1s
    input_series:
      - series: 'kube_node_status_condition{condition="Ready", job="kube-state-metrics", status="true", node="host1.lab.comp.com"}'
        values: '1+0x05'
      - series: 'ceph_disk_occupation{job="rook-ceph-mgr", exported_instance="host1.lab.comp.com"}'
        values: '1+0x05'
    promql_expr_test:
      - expr: cluster:ceph_node_down:join_kube
        eval_time: 4s
        exp_samples:
          - labels: 'cluster:ceph_node_down:join_kube{node="host1.lab.comp.com"}'
            value: 1
  - interval: 1s
    input_series:
      - series: 'kube_node_status_condition{condition="Ready", job="kube-state-metrics", status="true", node="host1.lab.comp.com"}'
        values: '0+0x05'
      - series: 'ceph_disk_occupation{job="rook-ceph-mgr", exported_instance="host1.lab.comp.com"}'
        values: '1+0x05'
    promql_expr_test:
      - expr: cluster:ceph_node_down:join_kube
        eval_time: 4s
        exp_samples:
          - labels: 'cluster:ceph_node_down:join_kube{node="host1.lab.comp.com"}'
            value: 0
